```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# Lab 04: Cracking the Vigenere Cipher

In today's lab, you'll learn how to crack the Vigenere cipher using the Index of Coincidence

Let's get started!

Run the cell below to import a function that you've written in the past:
* `text_clean`

```python
from lab04toolkit import text_clean
```

You can now use these functions in this notebook as if you'd written them here in the notebook:

```python
print( text_clean('Hello World!') )
```

Throughout this lab we'll be attempting to crack the following ciphertext. Run the cell to load it as a string assigned to the variable `ciphertext`.

```python
ciphertext = 'CSMULERWDLKAUVZRIASYLBMITFXAYMUSHZMLAQEKDFLWYAZWVLVGUBYIDHRCHVUSHOENPVXRQALAUOKSFVSFJMFVVDMULAYIJHHHLMGIFPRLVBYIDVSCOMIWKZXWYERWTLEVPVXFWAMLOIURQWMUACIIUVVUVVMITZELPWEWKUMLHVUAJHXAZBYIWZIGMISSQRXZVCXLVHPAJMNMVOSMAXZGVBVWZWIGQUZWYARXKVRKZWJLGDEKJWEWKKIJPVXMPOIJVEEQKUHSZEVPNHWKOMTSWSHXVZKLGOSLKIPQCKIZLZWIGSZWYGJPGLTQHVUWVBTAKEYIVOIJAPVTNLEKBZVSHTECPVXEFHMKFKYEKUAGBTUFGDSJAPKLGAVGBJCIQMKWABZRIBTSULGMERMFNBYIFHMKPMJAJLRKBLUIPSCSDPZXGYETIQKAKALHPVBIALWJHVTPQZITFPVVVOIJLERWPVXZPVXWQCIJFZVQCYOSITVMPALSAVFVFPHSSQTIVOMFRQKWQCIJFULGJVYLVNKLGDEQAWYICYXZLZRFDPXKHGKSKAWWSNFLFLEJVPUICYMKOICPDLPSAMNLGUWZLBYSWNLLPBFZGYEXAMIACYHKPBFGEBVJLLKSJLVLOIKWJLSMNPKXQOENLEFRFLVWKIKXJPWTBBRXVOILPUVMVHPDZMVQGKUMPBVRCAYJHTSYVDLWUBYITHFTPBRGVBEDSGKSQREOHBTLQBXGMQKWYHMKAKFEVWSURMKEPKPGVSVHCAMLHVUXJLRZBZIMGKSFHTZGGZXSYBVHVVLWYNVIVMSJPBWPCZLWKITVQZWZLZDMPKXZHBJLGOEVUMMITIIXVZVWGLRSYISFKAAAAPVMVOIJHERMUAGGHBGSERILVZRACAGZAWKEMLSMAWWMVHRVICIRKUKOPBYGWYMGZQKCUOIJHVRGTVWKAPVJKLPVHNKITPXSULWSTAYFHBVPADEKQCJXKUXATMKSULIAAXFTFVAFHTRVILVSIJZXJVPWBVUITALWOMUKG'
print(ciphertext)
```

## Question 1: The Index of Coincidence

The index of coincidence is a numerical statistic that indicates how likely it is that you select two letters from your ciphertext and those two letters are the same (e.g. both A's, or both B's, or both C's, etc.). Write a function that calculates this score.

When computing the probability of selecting 2 of a particular character we will use the *exact* calculation of:

$$\frac{n_A}{n} \cdot \frac{n_{A - 1}}{n - 1} + \frac{n_B}{n} \cdot \frac{n_{B - 1}}{n - 1} + \frac{n_C}{n} \cdot \frac{n_{C - 1}}{n - 1} + ... + \frac{n_Z}{n} \cdot \frac{n_{Z - 1}}{n - 1}$$

where $n_A$ is the number of A's in a message of length $n$ instead of the approximate calculation of: 

$$p_A ^ 2 + p_B^2 + p_C^2 + ... + p_Z^2$$

where $p_A$ is the proportion of A's in the message. 

For longer messages the difference in the calculation methods are fairly negligible, but if there is an error it will always be because the simplified calculation is a bit larger than the exact calculation. In this lab we're working with fairly short messages, hence we are using the **exact IC**, not the approximation.

```python
def index_of_coincidence(message, LETTERS='ABCDEFGHIJKLMNOPQRSTUVWXYZ'):
    cleaned_message = text_clean(message, LETTERS)
    total = 0
    message_length = len(cleaned_message)
    # YOUR CODE GOES BELOW THIS LINE
    

index_of_coincidence(ciphertext)
```

```python
grader.check("q1")
```

## Question 2: Guess the Cipher Type

Recall that the index of coincidence changes based on the nature of the text. The English language, or a ciphertext produced from a monoalphabetic cipher, should have an index value of around 0.068. A *perfect* polyalphabetic cipher, one that results in every ciphertext character being in equal proportion, would have an index value of 0.038. When computing the index of coincidence for a message, you're attempting to learn more about which type of cipher most likely produced the message.

For each of the following ciphertexts, guess the type of cipher that was used to create the message based on it's index of coincidence. Assign a value of 1 for a monoalphabetic cipher, and a value of 2 for a polyalphabetic cipher.

```python
print('Sample Ciphertext #1', index_of_coincidence("Dolu pu aol Jvbyzl vm obthu lcluaz, pa iljvtlz uljlzzhyf mvy vul wlvwsl av kpzzvscl aol wvspapjhs ihukz dopjo ohcl jvuuljalk aolt dpao huvaoly, huk av hzzbtl htvun aol wvdlyz vm aol lhyao, aol zlwhyhal huk lxbhs zahapvu av dopjo aol Shdz vm Uhabyl huk vm Uhabyl'z Nvk luapasl aolt, h kljlua ylzwlja av aol vwpupvuz vm thurpuk ylxbpylz aoha aolf zovbsk kljshyl aol jhbzlz dopjo ptwls aolt av aol zlwhyhapvu."))
print('Sample Ciphertext #2', index_of_coincidence('ZLGYIETAMQBXVUPOWHNUOAHZGYTJIMJSPRQGDNVCXAGNUCHZRFNXXSBSPGEOUILACYYIVSEGOEQHVFENMAEDLEVVFLJLVVCHVBRFXGOTYEFEWGKEPZTYEKIBQWSCDSLMXIABQKVSEGOPMFFRJVSEVAKBVGKIUPPRRTBSNQHGBURLLBOGLSPEONHBKVGKINLWJOYVOGXVGLNUOYVOGXVGDGFDXVHVWPGEHVMTLSPHRVCEJPXKHGRXJPOGIGQCAVSHXAEKBVREHUWTRVSMPOGWLGJSYONTRQHGNLRVTAMQNXWGDWYIVPWZSINEHVMMWHUHWGAAIAMQCA'))
print('Sample Ciphertext #3', index_of_coincidence("Tdlr vv gal Gqiljw st mbgmn rzzrgl, ab ufgqaqw frggwksrp dtf fbr ttsdap xh rlakgzgz xal tdztbbkcl mbnqv odpkj ohvz gqbargvxh walq ieba hnbhalv, rnq wh oskmgq emabt zal tdkavj gt yal ierka, aal wwtprrtx enq hukul dlttbwb gh kdpkj aal Plwo gt Sntnlv enq rt Sntnlv'w Yur hrgbbep xalq, m dhggrg kvwhtgv mh hal sdxvvwbf gt rmnxsvq uvukczvw laht malc qzvifo ghgnlrv xal gcumww odpkj pubtp ealq fh hal wwtprrtbwb."))
print('Sample Ciphertext #4', index_of_coincidence("Jr ubyq gurfr gehguf gb or frys-rivqrag, gung nyy zra ner perngrq rdhny, gung gurl ner raqbjrq ol gurve Perngbe jvgu pregnva hanyvranoyr Evtugf, gung nzbat gurfr ner Yvsr, Yvoregl naq gur chefhvg bs Unccvarff.--Gung gb frpher gurfr evtugf, Tbireazragf ner vafgvghgrq nzbat Zra, qrevivat gurve whfg cbjref sebz gur pbafrag bs gur tbirearq, --Gung jurarire nal Sbez bs Tbireazrag orpbzrf qrfgehpgvir bs gurfr raqf, vg vf gur Evtug bs gur Crbcyr gb nygre be gb nobyvfu vg, naq gb vafgvghgr arj Tbireazrag, ynlvat vgf sbhaqngvba ba fhpu cevapvcyrf naq betnavmvat vgf cbjref va fhpu sbez, nf gb gurz funyy frrz zbfg yvxryl gb rssrpg gurve Fnsrgl naq Unccvarff."))
```

```python
# For each ciphertext, assign a value of 1 for a monoalphabetic cipher, 
# and a value of 2 for a polyalphabetic cipher.
ciphertext_1 = ...
ciphertext_2 = ...
ciphertext_3 = ...
ciphertext_4 = ...
```

```python
grader.check("q2")
```

We'll assume that the `ciphertext` we loaded in at the start of this Lab was created with a Vigenere cipher (polyalphabetic) to complete the rest of the lab.

## Question 3: Splitting Text

In order to use the Index of Coincidence to crack the Vigenere cipher you will need to be able to work with particular subgroups of the ciphertext, where each group is created by collecting characters that would have been enciphered using the same letter in the keyword. To do this, you should write a function `text_group` to handle the dirty work.

The function should take in a ciphertext message, `message` and an integer `n`, and return back a list of strings, one for each group created assuming a key of length `n`. Make sure you clean the message before breaking it down into the groups.

**Hint:** String slicing, e.g. `message[a:b:c]`, could make this really simple! Think about the starting indices and the step size you'd want to use, and how they relate to the parameter `n`.

```python
def text_group(message, n):
    cleaned_message = text_clean(message)
    groups = []
    # YOUR CODE GOES BELOW THIS LINE
    

# Sample use of the function
text_group(ciphertext, 7)
```

```python
grader.check("q3")
```

## Question 4: Average Index of Coincidence

Once you have a ciphertext message that you believe to be enciphered using the Vigenere cipher, you'll know you guessed the keyword length correctly if the Index of Coincidence for each grouping of text that's created is approximately the same as English. This occurs because that means you successfully put all the characters from the ciphertext together with the other characters that were enciphered by the same letter in the keyword. Essentially all the characters in each group underwent the same Caesar cipher, so they should "look and feel" like characters you obtained from a Caesar cipher, including their Index of Coincidence.

Write a function `average_IC` that takes in a list of strings `groups`, where the list contains grouped text as determined by the `text_group function`. The function should compute the index of coincidence for each string, and then returns a float that represents the average Index of Coincidence of all the strings in the list. Your code should be flexible enough that you don't need to know how many strings are in the list ahead of time, it should work equally well for 2 groups or 200 groups!

```python
def average_IC(groups):
    total = 0
    # YOUR CODE GOES BELOW THIS LINE
    

# Sample use of the function
average_IC(text_group(ciphertext, 7))
```

```python
grader.check("q4")
```

## Question 5: Putting it All Together

Write a function, `get_keyword_length` that takes in a ciphertext, `message`, and then uses the functions `text_group` and `average_IC` to return the most likely length of the keyword. You should try several key lengths, starting with length 2, and increment the length by 1 until you compute an average index of coincidence that seems likely. Your function should give up if it doesn't find a likely key that is smaller than 10. If it doesn't find a key, it should return the integer -1.

```python
def get_keyword_length(message):
    cleaned_message = text_clean(message)
    # YOUR CODE GOES BELOW THIS LINE
    

# Sample use of the function
get_keyword_length(ciphertext)
```

```python
grader.check("q5")
```

You should now have a good idea about how long the keyword was that enciphered `ciphertext`. While not part of this lab, at this point you should easily be able to create a bar chart for each group of characters to determine the letter of the keyword that enciphered the characters in that group. Once you have all the letters in the keyword, you should easily be able to decipher the message!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 2)
```
>>> index_of_coincidence('Hello', LETTERS='ABCDEFGHIJKLMNOPQRSTUVWXYZ')
0.1
```
```
>>> round(index_of_coincidence('This is not a very long sentence but I hope that it will be long enough that the index score is about what you would expect from the english language or a mono alphabetic cipher'), 5)
0.06116
```
```
>>> round(index_of_coincidence('CSMULERWDLKAUVZRIASYLBMITFXAYMUSHZMLAQEKDFLWYAZWVLVGUBYIDHRCHVUSHOENPVXRQALAUOKSFVSFJMFVVDMULAYIJHHHLMGIFPRLVBYIDVSCOMIWKZXWYERWTLEVPVXFWAMLOIURQWMUACIIUVVUVVMITZELPWEWKUMLHVUAJHXAZBYIWZIGMISSQRXZVCXLVHPAJMNMVOSMAXZGVBVWZWIGQUZWYARXKVRKZWJLGDEKJWEWKKIJPVXMPOIJVEEQKUHSZEVPNHWKOMTSWSHXVZKLGOSLKIPQCKIZLZWIGSZWYGJPGLTQHVUWVBTAKEYIVOIJAPVTNLEKBZVSHTECPVXEFHMKFKYEKUAGBTUFGDSJAPKLGAVGBJCIQMKWABZRIBTSULGMERMFNBYIFHMKPMJAJLRKBLUIPSCSDPZXGYETIQKAKALHPVBIALWJHVTPQZITFPVVVOIJLERWPVXZPVXWQCIJFZVQCYOSITVMPALSAVFVFPHSSQTIVOMFRQKWQCIJFULGJVYLVNKLGDEQAWYICYXZLZRFDPXKHGKSKAWWSNFLFLEJVPUICYMKOICPDLPSAMNLGUWZLBYSWNLLPBFZGYEXAMIACYHKPBFGEBVJLLKSJLVLOIKWJLSMNPKXQOENLEFRFLVWKIKXJPWTBBRXVOILPUVMVHPDZMVQGKUMPBVRCAYJHTSYVDLWUBYITHFTPBRGVBEDSGKSQREOHBTLQBXGMQKWYHMKAKFEVWSURMKEPKPGVSVHCAMLHVUXJLRZBZIMGKSFHTZGGZXSYBVHVVLWYNVIVMSJPBWPCZLWKITVQZWZLZDMPKXZHBJLGOEVUMMITIIXVZVWGLRSYISFKAAAAPVMVOIJHERMUAGGHBGSERILVZRACAGZAWKEMLSMAWWMVHRVICIRKUKOPBYGWYMGZQKCUOIJHVRGTVWKAPVJKLPVHNKITPXSULWSTAYFHBVPADEKQCJXKUXATMKSULIAAXFTFVAFHTRVILVSIJZXJVPWBVUITALWOMUKG'), 5)
0.04374
```
```
>>> round(index_of_coincidence('IODKNNSFXUASXHROHAJGDSRSHNKMWFELFIRJOZHIAFKTEAHFCOAWZKCHIYTNJRREVNKUXUHODKRAMQWCHSZXLKTESSPPDSJJVMMEWGIIRPAAYGDOZHJDWFIUWBPWBBKTMPDGTIJWGDEGDSDZNEBRMSPSQRLAODSYZXAANOJQVLPOXPGDWZEAZPDSJJOZHEAAURQARARRKTMGBHHRSEOVQRRRSDWRAAQHGKYZSJSVMXGKRAAVPVTIEHCZKYEANWJDWOLJAFQZRNMYYPDWZLRNKMCFDSTEQWRQGVZSPQBQHTEPKAUGNHBAWRWBPWUWFBKEWMQCROKTMPDOBTRWFQHGKGQIRRSDCGDWZKNJRIIEAPKXHNBEJVAFOISQBZCBNHTSHCVFJHHVQVYKBSXUEQWLNEFIEFDSDSAAPQEHPMNYGEHIEFQGGEYHMNYAZZQHVJHAEAAHFSOACGXBBVQVJWMDSHJRELBQZPIEOVMHWKPUKUWBPWNJRRIRPORPLWKMCYKCWXBDSDGYKHTIFWBPXUAIZGBITAVGWPXINLDQEEWBOIBBOSMEHKTSJWGDECERXCFDCAXVJUGTVJHAEJKAMRNJRPMQJHXMXAWFIYENMFRPVAVOAHTEFAJQVLKBQGNHZQHUAFIEFWFAWLOAASGDVMMEARNVVCVFILARSMEHCRXUEFFIRJKUXUWGTCZWBZIEWHUQVZJAMPAOZHNLSMGRBIXIKLFQWFECZAUEQTANOGQPQKAPMFPIDFRZVQVSWHTIEYOXPRZVQVYEHFPRIWEWGNOZUHEZUXLWBPXUABMQROIUXRZVQVRTQQPYABFPLBCDWUAGQIZARFSYEJQMAWVMTCUKAVYZCRLRNCIRBJZKZRJHGVVJUAYGPCYIRPHTISAKILBIGTIGNIEXRZOZHYKJQHNIMFLBQUTXUAMAYACSEXJWGMQBOHUQCKFFEAPDQVFKBURUAFAAAKDURVKBMXYAOEXNNSSYYWFERBSAMMQABIMGDPXYRAMQWNJRKIYHCILNEFOYEHWZKBJVQVFDCGPQAFETNHSMRQOZQRQAFMRQWZIELOQMVEUWZKUAFEIYBZUORWMAYACZMHLIWZHSQZAJUAFYEAJSDWJDOFXUAQTEEWQFIEOCRXUATAYEOWEXRNGIIEAKQAVHZXINRSFSOATAYAZCGX'), 5)
0.0411
```


Tests for q2 (points: 2)
```
>>> ciphertext_1 in [1, 2] and ciphertext_2 in [1, 2] and ciphertext_3 in [1, 2] and ciphertext_4 in [1, 2]
True
```


Tests for q3 (points: 1)
```
>>> text_group('hello there', 5)
['HT', 'EH', 'LE', 'LR', 'OE']
```
```
>>> text_group('hi', 3)
['H', 'I', '']
```
```
>>> text_group('whoa', 3)
['WA', 'H', 'O']
```
```
>>> text_group('This is a longer message, did you group it correctly?', 8)
['TOSOTL', 'HNAUCY', 'IGGGO', 'SEERR', 'IRDOR', 'SMIUE', 'AEDPC', 'LSYIT']
```


Tests for q4 (points: 1)
```
>>> average_IC(['TOSOTL', 'HNAUCY', 'IGGGO', 'SEERR', 'IRDOR', 'SMIUE', 'AEDPC', 'LSYIT'])
0.09166666666666666
```
```
>>> average_IC(['QEANEGLEXOAWYTAVCEMEWMRXLNHILLXSGPTWIIAUBFMHUBVUHELXHECHFNESXOKOEEHTWIEKZHWLHZMPJNAVKAKUIFIHUNESMFSOEFINLLSGYTLYHEXCFLSWNOZIIAHOEHIVYUJMHLBYZODWFRLYAKMHUCPDBMGBSHLTG',
...  'ZHNZWUKIHEHVFMTWIEFWSSSOGYEHYOAMLWBPDNCDUDOCHSYOFWELYWLUOZLNLUSEUOYBPOLOLYALNEWUAWHEMFTLITOCHZFWGHTJUDXUUWADFCLLYAWYTWJCLNEHHOKZACHRICELUTWMODXCSYHUOEOCHEYTWNTWYASCN'])
0.048300073909829995
```
```
>>> average_IC(['QHNWLIOHYMVIMWMSLYIYXMPBINUUMCBYHWXYCUNLXUOUHBILZYLNMUNHKFUIICNFMHOUIULAYCYYXYLJNNIHHZHHYCMUBMOXFYYOMCCYBNBYLC',
...  'ZAZGKXEWFAWEFWSXGHHLAGWWDADFOUSUFLLELFZSLKEEYWOKLWLZWJWVMKLFOUZSGSJFXLWGFLLEWFWWLZHAKECVIJLLWZDWSLUKOUEDWGWHSG',
...  'ENEUEHAVTTCEESRONELOSLTPICBDHHVOEEHWHOENOSEOTPEOHAHEPAAEATITHHEWFTEDNUSDTLHACTSCOEIOOAIRUEHTYODCRHAEHHPTMTSATN'])
0.07311648596052266
```


Tests for q5 (points: 2)
```
>>> get_keyword_length('CSMULERWDLKAUVZRIASYLBMITFXAYMUSHZMLAQEKDFLWYAZWVLVGUBYIDHRCHVUSHOENPVXRQALAUOKSFVSFJMFVVDMULAYIJHHHLMGIFPRLVBYIDVSCOMIWKZXWYERWTLEVPVXFWAMLOIURQWMUACIIUVVUVVMITZELPWEWKUMLHVUAJHXAZBYIWZIGMISSQRXZVCXLVHPAJMNMVOSMAXZGVBVWZWIGQUZWYARXKVRKZWJLGDEKJWEWKKIJPVXMPOIJVEEQKUHSZEVPNHWKOMTSWSHXVZKLGOSLKIPQCKIZLZWIGSZWYGJPGLTQHVUWVBTAKEYIVOIJAPVTNLEKBZVSHTECPVXEFHMKFKYEKUAGBTUFGDSJAPKLGAVGBJCIQMKWABZRIBTSULGMERMFNBYIFHMKPMJAJLRKBLUIPSCSDPZXGYETIQKAKALHPVBIALWJHVTPQZITFPVVVOIJLERWPVXZPVXWQCIJFZVQCYOSITVMPALSAVFVFPHSSQTIVOMFRQKWQCIJFULGJVYLVNKLGDEQAWYICYXZLZRFDPXKHGKSKAWWSNFLFLEJVPUICYMKOICPDLPSAMNLGUWZLBYSWNLLPBFZGYEXAMIACYHKPBFGEBVJLLKSJLVLOIKWJLSMNPKXQOENLEFRFLVWKIKXJPWTBBRXVOILPUVMVHPDZMVQGKUMPBVRCAYJHTSYVDLWUBYITHFTPBRGVBEDSGKSQREOHBTLQBXGMQKWYHMKAKFEVWSURMKEPKPGVSVHCAMLHVUXJLRZBZIMGKSFHTZGGZXSYBVHVVLWYNVIVMSJPBWPCZLWKITVQZWZLZDMPKXZHBJLGOEVUMMITIIXVZVWGLRSYISFKAAAAPVMVOIJHERMUAGGHBGSERILVZRACAGZAWKEMLSMAWWMVHRVICIRKUKOPBYGWYMGZQKCUOIJHVRGTVWKAPVJKLPVHNKITPXSULWSTAYFHBVPADEKQCJXKUXATMKSULIAAXFTFVAFHTRVILVSIJZXJVPWBVUITALWOMUKG')
8
```
```
>>> get_keyword_length('IODKNNSFXUASXHROHAJGDSRSHNKMWFELFIRJOZHIAFKTEAHFCOAWZKCHIYTNJRREVNKUXUHODKRAMQWCHSZXLKTESSPPDSJJVMMEWGIIRPAAYGDOZHJDWFIUWBPWBBKTMPDGTIJWGDEGDSDZNEBRMSPSQRLAODSYZXAANOJQVLPOXPGDWZEAZPDSJJOZHEAAURQARARRKTMGBHHRSEOVQRRRSDWRAAQHGKYZSJSVMXGKRAAVPVTIEHCZKYEANWJDWOLJAFQZRNMYYPDWZLRNKMCFDSTEQWRQGVZSPQBQHTEPKAUGNHBAWRWBPWUWFBKEWMQCROKTMPDOBTRWFQHGKGQIRRSDCGDWZKNJRIIEAPKXHNBEJVAFOISQBZCBNHTSHCVFJHHVQVYKBSXUEQWLNEFIEFDSDSAAPQEHPMNYGEHIEFQGGEYHMNYAZZQHVJHAEAAHFSOACGXBBVQVJWMDSHJRELBQZPIEOVMHWKPUKUWBPWNJRRIRPORPLWKMCYKCWXBDSDGYKHTIFWBPXUAIZGBITAVGWPXINLDQEEWBOIBBOSMEHKTSJWGDECERXCFDCAXVJUGTVJHAEJKAMRNJRPMQJHXMXAWFIYENMFRPVAVOAHTEFAJQVLKBQGNHZQHUAFIEFWFAWLOAASGDVMMEARNVVCVFILARSMEHCRXUEFFIRJKUXUWGTCZWBZIEWHUQVZJAMPAOZHNLSMGRBIXIKLFQWFECZAUEQTANOGQPQKAPMFPIDFRZVQVSWHTIEYOXPRZVQVYEHFPRIWEWGNOZUHEZUXLWBPXUABMQROIUXRZVQVRTQQPYABFPLBCDWUAGQIZARFSYEJQMAWVMTCUKAVYZCRLRNCIRBJZKZRJHGVVJUAYGPCYIRPHTISAKILBIGTIGNIEXRZOZHYKJQHNIMFLBQUTXUAMAYACSEXJWGMQBOHUQCKFFEAPDQVFKBURUAFAAAKDURVKBMXYAOEXNNSSYYWFERBSAMMQABIMGDPXYRAMQWNJRKIYHCILNEFOYEHWZKBJVQVFDCGPQAFETNHSMRQOZQRQAFMRQWZIELOQMVEUWZKUAFEIYBZUORWMAYACZMHLIWZHSQZAJUAFYEAJSDWJDOFXUAQTEEWQFIEOCRXUATAYEOWEXRNGIIEAKQAVHZXINRSFSOATAYAZCGX')
5
```
```
>>> get_keyword_length('QLLVCEARNGYLAWVSWBYTIHVCOWXCZCFVGFGIZVYTLCBEBPOSISEYTLWJCLDXHSJBSCZVNLLXICPDRGKCKEEHJQLPKBYHFNGYEWITXVXXOMGNPDQEXFCLVVRHHAWUJZEVIKDXRXXYYTVEYIZYOXRVNEIXYNLLAYGHBJXORWHVULUPZBUKPYWXYNQMTLKBISIQJVPWEDLAIRRUIJUINLYOBKYUNRKMKFIAPYDHNNGOECVVWGYGABIKOOXTSERSEWFZQHVEZUNAIALRVVWKBEABBVFCQSECMHVGCUVLBBVJKDWOJFLZGYCQWMFKOOVXYKQIJIGUVHBCFU')
-1
```
