```python
# Initialize Otter
import otter
grader = otter.Notebook("activity-13-loops-in-parallel.ipynb")
```

# Activity 13: Looping over Strings in Parallel

In **Activity 13: The Tabula Recta**, you saw how you could use a collection of characters as a key to work with the Tabula Recta. This is different than the keys you've work with so far, which have been an integer (or integers). As a result, you'll need to determine how to start using strings as keys alongside strings as messages.

## Step 1: Looping in Parallel

Suppose you have two lists of equal length. One contains student names and the other contains their grade on a recent Quick Check quiz.

```python
names = ['Alexis', 'Dee', 'Daryl', 'Sebastian', 'Ora']
grades = [2, 3, 2, 3, 1]
```

Suppose you want to generate a report that prints out a student's name with their grade next to it. If you were trying to avoid a loop, you could do that the **long** way by typing:

```python
print(names[0], grades[0])
print(names[1], grades[1])
print(names[2], grades[2])
print(names[3], grades[3])
print(names[4], grades[4])
```

But that seems a bit repetitive! (Also, did you see how if you give the `print` function 2 arguments it will print them side by side, with just a space in between, EVEN if they are different data types? Nice!)

## Question 1.1
Write a `for` loop below that results in the same output printed to the screen.

**Hint:** Thank about how you can compute how many items are in the lists, and then how to use that value to create a `range` of values that would correspond to the indices. Since the lists are the same length, any index that works for one is guaranteed to work on the other.

<!-- BEGIN QUESTION -->



```python
...
    ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.2

Strings and lists aren't that different as they each are a collection of smaller objects, but a string is a collection specifically of individual characters. In the last question you looped over two lists of equal length and printed the element found at each index. 

**Do the same for the two strings of equal length below.** The output should look like:
```
s b
e l
c a
r i
e s
t e
```

Remember, you should be writing a loop that references the strings by name and should work no matter what characters the strings contain. (e.g. don't write code that would only work with the provided strings, it should work without modification no matter their contents)

```python
q12_keyword = 'secret'
q12_message = 'blaise'

...
    ...
```

<!-- END QUESTION -->

#### Using Indices With Strings

You likely have already realized that this question is easiest solved by not iterating character by character over each string, but rather index value by index value. By writing a loop that starts:

`for i in range(len(message))`

You can use `i`, the string index to reference corresponding characters in *any* string, in this case both `keyword` and `message`. If you had started with a loop like:

`for char in message:`

you would have had an easy way to pull each character out of the `message` string, but no easy way to match it up with the corresponding character in `keyword`. For parallel looping, indices are your friend!

### Strings of Different Lengths

So, what can you do when strings have unequal length? The Tabula Recta method assumes that your keyword is equally long as the message, so each character in the message has a corresponding character in the key to work with. Keys are usually decided upon ahead of time, but the length of messages will vary from message to message. We'll need a way to make sure that any given key is properly used with variable length messages. We saw in the lesson that you could just repeat the key until it's long enough to "cover" the message. To do that in Python we have two ways to mimic this step:

1. Take your key as a string and repeatedly append copies of itself to the end until it's long enough to match the message.
2. Use a mod operator `%` to ensure that the correct index from the key is referenced , even when index values for the message get to be larger than the index values for the key.

Let's explore both options so you can see which you prefer.


### Question 1.3

In the cell below, compute the length of both the key and message, and then use that information to compute how many times the keyword would need to be repeated to "cover" the message. You should *not* manually count the length of each string, use Python code to do this.

```python
q13_key = 'ABCDEFG'
q13_message = 'thisisalongermessage'

q13_key_length = ...
q13_message_length = ...

times_to_repeat = ...
long_key = q13_key * times_to_repeat

print(times_to_repeat)
print(long_key)
print(q13_message)
```

```python
grader.check("q1_3")
```

The code cell above used a nice trick using strings. You can use the multiplication operator `*` to create a repeated copy of itself that was saved to `long_key`. Hopefully you saw that this multiplied version is now as long or longer than the string contained in `message`. You should now be able to run the cell below:

```python
for i in range(len(q13_message)):
    print(long_key[i], q13_message[i])
```

<!-- BEGIN QUESTION -->

## Question 1.4

The other option to work with keys and messages of unequal length is to use the mod operator `%` combined with the index for the message to compute the index of the corresponding character in the key. For example, for the character at index `107` in your message, which character in your `8` character key should be used to encipher the message? You should be able to compute that value using the current index `107`, and the length of your key, `8`.

Once you determine how to complete that computation, finish the following loop to print the correct character pairings between the two strings. The output should read:

```
A h
B e
C l
A l
B o
C f
A r
B i
C e
A n
B d
```

```python
q14_key = 'ABC'
q14_message = 'hellofriend'

for i in range(len(q14_message)):
    ...
```

<!-- END QUESTION -->



Tests for q1_3 (points: 1)
```
>>> len(long_key) >= len(q13_message)
True
```
```
>>> times_to_repeat
3
```
