```python
# Initialize Otter
import otter
grader = otter.Notebook("activity-22-prng-stream-cipher.ipynb")
```

# Activity 22: Creating a Pseudorandom Number Generator Stream Cipher

In **Activity 22: Introduction to Stream Ciphers**, you saw how you could use ASCII to convert a keyword and plaintext into a stream of binary that could be XOR'd to produce binary ciphertext. In an ideal environment, the keystream would not be based on a word that could be guessed, but rather contain a truly random sequence of 1's and 0's. Since true randomness is essentially impossible to capture, and even if it was, it would be actually impossible to ensure that both the sending and receiving parties had access to the same stream to encrypt and decrypt the message. After all, if you had a way to securely send over the random stream of 1's and 0's that was unique to each message so the receiver could use it, then why bother with encryption in the first place? Just send the message in the same fashion that you are sending the keystream! 

We can attempt to mimic a true One-time Pad (OTP) Cipher using the pseudorandom number generator that Python has built-in and is found in the `random` library.

## Question 1.1: Choose a message

Store an ASCII compliant message as a string to the variable `ascii_plaintext` below.

```python
ascii_plaintext = ...
```

```python
grader.check("q1_1")
```

## Question 1.2: Convert plaintext to binary

Write a `for` loop that iterates over `ascii_plaintext` one character at a time and uses `ord` to convert the character to it's ASCII/UTF-8 decimal value. Then, use the `format` function to convert this into an 8-bit string. Finally, append these 8 bits (1 byte) to the end of the string named `binary_plaintext`.

```python
binary_plaintext = ''
...
    ...

print(binary_plaintext)
```

```python
grader.check("q1_2")
```

## Question 1.3: Generate a keystream
You'll need to generate a keystream that has the same number of bits that your binary plaintext uses. 

You can do this by:
* set a seed of your choosing to initialize the random number generator. The notebook starts with a seed value of 4200 (the course number) but you can change it to anything you want. This value should be kept secret from prying eyes!
* using the `random.getrandbits()` function. Remember this function takes in an argument that specifies how many random bits you want. **Hint:** the length of the `binary_plaintext` string could be helpful here!
  * Remember, `.getrandbits()` will return an integer in decimal form by default. **That's okay**. We can use it to XOR in this form, even though we think about XOR using binary number. Python can convert for us behind the scenes as needed.

```python
import random
random.seed(4200)
keystream = ...
print(keystream)
```

```python
grader.check("q1_3")
```

### Question 1.4: XOR the plaintext and keystream
Now that you have plaintext as a binary string and the keystream as a decimal, you are just about ready to encrypt your message. First, convert `binary_plaintext` to it's decimal representation using the `int` function. Store the result to `numerical_plaintext`. Then use the XOR operation (`^`) in Python to XOR `numerical_plaintext` and `keystream` and save this result as `numerical_ciphertext`.

```python
numerical_plaintext = ...
numerical_ciphertext = ...
print(numerical_ciphertext)
```

```python
grader.check("q1_4")
```

### Question 1.5: Convert ciphertext to ASCII
Now that you have your ciphertext in decimal format, it's time to work it back to binary and then to ASCII so we can visually inspect it as a string.

Below is code that will convert your decimal to a binary representation of equal length as the original message. It does this by constructing the second argument of the `format` function using 3 concatenated strings: 
* `0`: the leading 0 indicates to pad this binary number using 0's
* `str(len(binary_plaintext))`: this computes the length of the plaintext (in binary) and then converts that number to a string
* `b`: the trailing `b` indicates to format the number in binary.

So if your original message was 40 binary digits, the second argument in the code below would be essentially the string `'040b'`

If you didn't include these specifics, and just included `'b'` as the second argument, you run the risk that the decimal representation of the ciphertext doesn't need as many bits as the plaintext did (maybe it had a few leading 0's) and therefore doesn't have a quantity of bits that is divisible by 8. This way, we can collect groups of 8 bits to convert to characters without worrying that we are "missing" some bits at the front of the binary representation.

```python
binary_ciphertext = format( numerical_ciphertext, '0'+str(len(binary_plaintext))+'b')
print(binary_ciphertext)
```

In the cell below, use a `loop` to iterate over 8 bits of `binary_ciphertext` at a time, each time converting those 8 bits to a decimal using `int`, and then using `chr` to convert the numerical value to the corresponding ASCII character. This character should be appended to the end of the string named `ascii_ciphertext`.

After conversion, look at how `print(ascii_ciphertext)` and `print(ascii(ascii_ciphertext))` differ. If there are any 8-bit numbers that correspond to a decimal value outside of the standard printable ASCII range of 32-127 then the hexadecimal code will be displayed instead. So characters that only appear in extended-ASCII or UTF-8, for example `Û`, will display like `\xdb` and other non-printable ASCII characters like `` will display as `\x11`. 

**Remember:** All the information is correctly stored in the string `ascii_ciphertext`, using the `ascii` function just decides how to display that information when printing it out. 

```python
ascii_ciphertext = ''

...
    ...

print(ascii_ciphertext)
print(ascii(ascii_ciphertext))
```

```python
grader.check("q1_5")
```

### Question 1.6: Deciphering the message

Run the cell below to  confirm that when you XOR the ciphertext with the keystream, the result is the original plaintext message. In this example it's sufficient to verify that their numerical representations are identical since we know that the numerical representations determine the text.

```python
verify_plaintext = numerical_ciphertext ^ keystream
print('             plaintext (decimal): ', numerical_plaintext)
print('ciphertext ^ keystream (decimal): ', verify_plaintext)
```

## Wrapping Up

That's it! You've found a way to generate a pseudorandom stream of bits that can be used with the XOR operation to implement a stream cipher in Python. As long as the person you send your message to knows the `seed` value you used to prepare the random number generator, they can decipher any message you send them. You should *of course* never reuse the same seed however, as reuse of the key can allow for an attacker to collect ciphertexts that were encrypted used the same keystream and eventually determine the keystream.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1_1 (points: 1)
```
>>> ascii_plaintext.isascii()
True
```


Tests for q1_2 (points: 1)
```
>>> len(binary_plaintext) == len(ascii_plaintext) * 8
True
```


Tests for q1_3 (points: 1)
```
>>> random.seed(4200)
>>> keystream == random.getrandbits(len(binary_plaintext))
True
```


Tests for q1_4 (points: 1)
```
>>> numerical_ciphertext == numerical_plaintext ^ keystream
True
```


Tests for q1_5 (points: 1)
```
>>> len(ascii_ciphertext) == len(ascii_plaintext)
True
```
