```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Homework 04: Fractionation in ADFGVX, Polybius Square, and Related Ciphers ¶

Today's easy access to computing power means that modern ciphers must be more sophisticated that pre-information age ciphers. Even in an LFSRsum, a single change in a character will change only one character in the ciphertext. Furthermore, this change in the plain text will align with the change in the ciphertext.

For example, look at these two messages enciphered with our LFSR from Lab06 last week. The key is 01011. Can you decipher them? Type the plaintext in the cell below.

YXOJZBJYN.
YXOJZBJIN.

<i>Recall that the "rule" for this LFSR is that b4 = b3 ⊕ b2 ⊕ b1</i>

_Type your answer here, replacing this text._

<!-- END QUESTION -->

This leaves this kind of crypto system vulnerable to a plaintext attack, if a message is known to begin with, end with, a particular word or phrase (and many might begin by addressing the person for whom the message was intended, which is a piece of information the intercpetor probably already knows).

The following message was encrypted in the same way as the two above and with the same "rule". The first word is the same. 



Can you recover the plaintext by working backwards? Type your answer into the cell below. 

RAPAOAAOLXKR


_Type your answer here, replacing this text._

<!-- END QUESTION -->

In block cipher, our message is broken into parts which are then altered, combined, and reshuffled to create our ciphertext. In preparation for <b>that</b> cipher, we will write a program for the Polybius cipher.

## Part I ##
The Polybius Cipher

Here is the Polybius square or Polybius checkerboard from our webinar followed by the decipher function:

```python
square = [['E','M','O','I','S'],['T','H','A','G','N'],['Z','D','W','Y','B'],['C','F','K','L','P'],['Q','R','U','V','X']]
```

```python
print(square[0][1])
```

```python
def polybius_decipher(message):
    square = [['E','M','O','I','S'],['T','H','A','G','N'],['Z','D','W','Y','B'],['C','F','K','L','P'],['Q','R','U','V','X']]
    output = ""
    for i in range(0,len(message),2):
        y = (message[i]+message[i+1])
        output += (square[int(y[0])-1][int(y[1])-1])
    return(output)    

```

```python
print(polybius_decipher("42441111"))
```

In order to work backwards, we need to input a character and get back the two digit location of the number. Notice that asking the program to "find" a character doesn't work. Remove the hashtag to see.

```python
#print(square.find("H"))
```

How can we type a message and encipher it correctly? One way is to <b>use a for loop inside of a for loop</b> to run through all the possible permutations of indices from 0 to 4 and examine each character for a match. When the program finds a match, it returns the two digit index value.

Note that you will have to use two different variables, one in each for loop. So if the first one is "for i in range..." the second will have to be "for j in range..."

Write your program in the cell below. Using our square and the character "P" should generate the digits 45.

If you get 34 instead, recall that although Python will index starting with zero, humans have tradtionally indexed starting with 1, so to replicate the pencil paper Polybius, you have to adjust this.

```python
def poly_find(char):
    ...
            
```

```python
grader.check("q2_1")
```

## Part II: Square Dancing

In our webinar, we used a key phrase to load fill in the Polybius square. Can you write a program that takes in a phrase and generates the square?

Before you begin, think about the strings we have used for letters. Here we use one that is different from the ones that have come before. How is it different and why? Type your answer into the cell below.

<!-- BEGIN QUESTION -->



```python

```

<!-- END QUESTION -->

Write a function that will generate a square with the phrase "What does the fox say?" The last line in this square should be "['K', 'L', 'M', 'N', 'O']"

```python
def square_gen(key, LETTERS = "ABCDEFGHIKLMNOPQRSTUVWXYZ"):
    key = key.upper()
    square = [[],[],[],[],[]]
...
    
```

```python
grader.check("q3_2")
```

## Part III: Underground Newspapers

It is April 9th, 1940. German forces have begun their occupation of Norway. King Haakon VII is in exile in England. You are Henrik Weber, editor of the underground newspaper <i>For konge og fedreland</i> ("For King and Fatherland"). In defiance of the Nazi occupiers, you risk your life to bring your people news from London. 

One of your contacts in the resistance knows of someone who can provide a safe place for the production of the paper. Use the word "Preikestolen" as a key to learn the name of the person who can help you and the location of his secret room.

2114414512344213123213122542312131111214252214254111121321211425312145434313124223452113
23252242131421243125342335151424241425411325


In the cell below, write a program that takes in an encrypted message and a key word or phrase and returns the plaintext.

<b><i>Hint!</i></b> You will have to revise the program we wrote in webinar (and which appears at the top of this notebook) so that it takes in a key.

```python
def polybius_decipher(message, key):
    square = square_gen(key, LETTERS = "ABCDEFGHIKLMNOPQRSTUVWXYZ")
    output = ""
...
    
```

```python
grader.check("q4_1")
```

Who is your contact and where is his secret room? Type your answer in the cell below.

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2_1 (points: 3)
```
>>> poly_find('H') != None
True
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> poly_find("V")
'54'
```
failure_message: Check that you are adjusting the indicies for each row and column to make the program behave more like a human


Tests for q3_2 (points: 2)
```
>>> square_gen("HelloWorld", LETTERS = "ABCDEFGHIKLMNOPQRSTUVWXYZ") != None
True
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> print(square_gen("HelloWorld", LETTERS = "ABCDEFGHIKLMNOPQRSTUVWXYZ"))
[['H', 'E', 'L', 'O', 'W'], ['R', 'D', 'A', 'B', 'C'], ['F', 'G', 'I', 'K', 'M'], ['N', 'P', 'Q', 'S', 'T'], ['U', 'V', 'X', 'Y', 'Z']]
```
failure_message: Check that you are adding unused letters to the end of your square
```
>>> print(square_gen("visualizewhirldpeas"))
[['V', 'I', 'S', 'U', 'A'], ['L', 'Z', 'E', 'W', 'H'], ['R', 'D', 'P', 'B', 'C'], ['F', 'G', 'K', 'M', 'N'], ['O', 'Q', 'T', 'X', 'Y']]
```


Tests for q4_1 (points: 2)
```
>>> polybius_decipher("1112131415", "abcde") != None
True
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> polybius_decipher("1112131415", "ABCDE")
'abcde'
```
failure_message: Check that you are returning plaintext in lower case.
