```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

# Lab 03: Creating a Toolkit of Functions

In today's lab, you'll learn how to:

1. program the Extended Euclidean algorithm to determine a multiplicative inverse in any modulus.
2. store the functions you've written into a personal "toolkit" that you can use in your coursework or other projects.

Let's get started!

## Part 1: Extended Euclidean Algorithm

Now that you know a little about lists, programming the Extended Euclidean algorithm should be a little bit easier to think about. Recall the example from an earlier lesson, where you were looking to find the multiplicative inverse of 707 in modulus 979. We set up the following table:

$$\begin{array}{c|rr}
     & 979 & 707 \\ \hline
 979 & 1 & 0 \\
 707 & 0 & 1 \\
 272 & 1 & -1 \\
 163 & -2 & 3 \\
 109 & 3 & -4 \\
  54 & -5 & 7 \\\
   1 & 13 & -18
\end{array}$$

This table allowed a structure in which we performed operations on the top two rows to generate information about a new third row. Then, we repeated the process with the second and third rows as our "new" first and second rows to generate a "new" third row. This process repeated until we obtained a row whose first element was 1. Let's remember what those operations were that we used to generate new rows in the table.

We started with the top two rows, which are always set up the same. The larger number goes in the first row, with a 1 and then a 0, and the smaller number went into the second row, followed by a 0 and then a 1.

$$\begin{array}{c|rr}
     & 979 & 707 \\ \hline
 979 & 1 & 0 \\
 707 & 0 & 1 
\end{array}$$

To obtain the next row in the table, we saw we could subtract 707 **once** from 979 (without obtaining a result less than our goal of 1).  We then also subtract 0 **once** from 1 (second column) and 1 **once** from 0 (third column) to obtain the third row of $272, 1,$ and $-1$. 

$$\begin{array}{c|rr}
     & 979 & 707 \\ \hline
 979 & 1 & 0 \\
 707 & 0 & 1 \\
 272 & 1 & -1
\end{array}$$

How did we know to only subtract off one multiple of the elements in row 2 from row 1? We could have visually inspected the numbers 707 and 979 to determine that, but it would be helpful for us to think about how to **compute** this value so we can instruct our computer to do the same operation. For example, if rows 1 and 2 had the numbers $13475942$ and $324235$, it would be a lot tricker to do this calculation by hand. You'd have to compute some multiples of $324235$ to see how many it takes to get the result larger than $13475942$. Not too *difficult* to do, but it could take awhile... and don't we have some mathematical operation that computes how many multiples of a number "go into" a larger number?

You may remember from an earlier math class learning about the long division algorithm. If you need a refresher, check out this link: https://www.mathsisfun.com/long_division.html

$$3{\overline{\smash{\big)}\,130\phantom{)}}}$$

Using the algorithm you could compute that 3 (the divisor) "goes into" 130 (the dividend), 43 times (the quotient) with a remainder of 1. It turns out, Python can compute each of those values for you using some operators, one of which you already know.

* The `//` (double division symbol) operator in Python will compute the quotient (`130 // 3`). In mathematics this is often represented with the help of the floor function $\lfloor \phantom{x} \rfloor$, for example $\lfloor 130\phantom{x}/\phantom{x}3 \rfloor$
* The `%` operator will compute for you the remainder, a fact that you may have already picked up on in our earlier work in the course. 

To see these operators in action, run the following cells.

```python
print( 130 // 3 )
```

```python
print( 130 % 3 )
```

The `//` operator gives us a way to quickly compute how many multiples of the second row should be subtracted off from the first row, even when the numbers are large!

```python
print( 13475942 // 324235 )
```

For the task at hand, let's take a look at how many times 272 "goes into" 707.

```python
print( 707 // 272 )
```

Now that we know it should be twice the number in the second row subtracted from the number in the first row, we can compute the third row.

$$\begin{array}{c|rr}
     & 979 & 707 \\ \hline
 707 & 0 & 1 \\
 272 & 1 & -1 \\
 163 = 707 - 2(272) & -2 = 0 - 2(1) & 3 = 1 - 2(-1)
\end{array}$$

We can now keep doing this process until we compute a row that starts with 1. The last few rows of this example are:

$$\begin{array}{c|rr}
     & 979 & 707 \\ \hline
 109 & 3 & -4 \\
  54 & -5 & 7 \\\
   1 = 109 - 2(54) & 13 = 3 - 2(-5) & -18 = -4 - 2(7)
\end{array}$$

Once we have the final row, we know that: $$1 = 979(13) + 707(-18)$$ which implies that $$1 \equiv 707(-18) \pmod{979}$$

We can quickly compute the multiplicative inverse of 707 by taking the last number in the last row, and modding by the modulus, 979.

```python
print( -18 % 979 )
```

Let's think about how we could generalize this process so we can write a function to do all this for us!

Suppose we start with two numbers, $n$ and $m$ where $m$ is assumed to be the larger of the two numbers. The goal is to find the multiplicative inverse of $n$ in modulus $m$. Recall that this means $1 \equiv n \cdot n^{-1} \pmod{m}$ 

We could set up row 1 and row 2 in the table:

$$\begin{array}{c|rr}
     & m & n \\ \hline
 m & 1 & 0 \\
 n & 0 & 1 
\end{array}$$

To compute row 3, we need to know how many multiples of $n$ (we'll call this value $k$) to subtract off from $m$. We can compute $k$ using the `//` operator, `k = m // n`, and then determine row 3 as:

$$\begin{array}{c|rr}
     & m & n \\ \hline
 m & 1 & 0 \\
 n & 0 & 1 \\
 m-kn & 1 - k(0) & 0 - k(1) \\
\end{array}$$

After you do this for a few rows, you'd end up with a table in the form of:

$$\begin{array}{c|rr}
     & m & n \\ \hline
 ... & ... & ... \\
 p & b_p & a_p \\
 q & b_q & a_q \\
 p - kq & b_p - kb_p & a_p - kq_q
\end{array}$$

If we ever get a row that starts with $1$, then the GCD between $m$ and $n$ is 1, which implies that $a$ (or $a \pmod{m})$ is a multiplicative inverse of $n$ modulo $m$.

### Question 1.1

We'll start off with two values for $m$ and $n$ such that $m > n$. 

Create the first two rows of the Extended Euclidean Algorithm table, where each row is represented at a list. The first element of the list should store the value from the first column of the table, the second element in the list should store the value from the third column of the table, and the third element of the list should store the value from the third column of the table.

```python
m = 1323
n = 220

row1 = ...
row2 = ...
```

```python
grader.check("q1_1")
```

### Question 1.2

Now, compute the number of multiples of row 2, $k$, that should be subtracted from row 1. You should reference the values at index 0 in each list (row 1 and row 2), not the values for $m$ and $n$ directly.

```python
k = ...
print( k )
```

```python
grader.check("q1_2")
```

### Question 1.3

Now that we've defined row 1 and row 2 as lists, and we know how many multiples of row 2 need to be subtracted off from row 1, let's define a list for row 3, based on the two lists that hold the information about the previous 2 rows. Your definition should take the form of:

`row3 = [ <computation for element 0>, <computation for element 1>, <computation for element 2> ]`

```python
row3 = ...
print( row3 )
```

```python
grader.check("q1_3")
```

### Question 1.4

We'll now want define the rows such that row 2 $\rightarrow$ row 1, and row 3 $\rightarrow$ row 2 so we can repeat the process over and over until the list we compute in row 3, has a 1 as the first element.

We're ready to take our work and move it into a function so we can try it out with different values for $n$ and $m$. As you fill in the function below, you should be able to reuse some of your code from earlier questions. The part you'll want to think carefully about is how to write a loop (what kind? when should it stop running?) so that when the function is done running, `row2` contains the list `[1, b, a]` where `b` and `a` are the final coefficients. 

**Remember:** The final `a` value in row 2 will be the multiplicative inverse of $n$ modulo $m$.

```python
def mult_inverse(n, m):
    row1 = ...
    row2 = ...
    
    # What kind of loop should go here, and under what condition should it stop running?
    ...
        k = ...

        row3 = ...

        row1 = ...
        row2 = ...
    
    inverse = ...
    
    return inverse
```

```python
grader.check("q1_4")
```

## Part 2: Building a Toolkit

Now that you have this function you like, you should plan to hold on to it for future use! You can copy it's code and put it into a Python file for reference later. This part will walk you through creating your own Python module that you can use in a similar manner to how we've used `matplotlib` to generate bar charts.

### Importing From a Python File
You'll notice there is a file named `lab03toolkit.py` in the same folder as this notebook. This file is called a Python file, and it contains ONLY Python code. It's different from a Python Notebook that we mostly work in, mostly because it can't hold any other text, images, links, etc AND when you run a Python file all the code runs at once. You can't break it down into smaller chunks of code.

Double clicking on this file will open it in a new tab. You'll notice there is code for the `text_clean` function that you were asked to write in previous activities. If you wanted to use this function in your notebook, you could do so by running the cell below:

```python
from lab03toolkit import text_clean
```

You can probably tell from the wording that this line of code will import the `text_clean` function from the file named `lab03toolkit`. (**Notice:** you don't include the `.py` in the import statement). This import method will only work if your toolkit file is in the same folder as your notebook! There *are* ways to pull in other Python files in different folders, but it's beyond the scope of this course.

Run the cell below to verify that it correct imported `text_clean`.

```python
print( text_clean('Messy T3xt!') )
```

Using a Python file as a toolkit to hold all your most-used functions allows you to quickly copy your toolkit into a folder with a notebook and immediately have access to your favorite tools! In this course, I will always provide the necessary functions for a Lab or Homework assignment in an accompanying toolkit file, but that won't always be the case for an activity.

### Question 2.1

Copy your `mult_inverse` function from this notebook, paste it into `lab03toolkit.py`, **rename the function to `multiplicative_inverse`**, and then save the file. In the cell below, import the newly renamed function from the toolkit, then run the tests to make sure you did this correctly.

```python
...
```

```python
grader.check("q2_1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, find the `.zip` file in the left side of the screen in the file browser, right-click, and select **Download**. You'll submit this `.zip` file for the assignment Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_1 (points: 1)
```
>>> m > n
True
```
```
>>> type(row1) == list
True
```
```
>>> type(row2) == list
True
```
```
>>> len(row1) == 3
True
```
```
>>> len(row2) == 3
True
```
```
>>> row1 == [m, 1, 0]
True
```
```
>>> row2 == [n, 0, 1]
True
```


Tests for q1_2 (points: 1)
```
>>> k
6
```


Tests for q1_3 (points: 1)
```
>>> len(row3) == 3
True
```
```
>>> row3
[3, 1, -6]
```


Tests for q1_4 (points: 2)
```
>>> mult_inverse(220, 1323)
439
```
failure_message: Your function incorrectly determined 220^-1 in modulo 1323
success_message: Your function correctly determined 220^-1 in modulo 1323
```
>>> mult_inverse(707, 979)
961
```
failure_message: Your function incorrectly determined 707^-1 in modulo 979
success_message: Your function correctly determined 707^-1 in modulo 979
```
>>> error = False
>>> try:
...     mult_inverse(6, 18)
... except ZeroDivisionError:
...     error = True
>>> error == True
True
```
failure_message: Your function should produce a ZeroDivisonError if the gcd(n, m) != 1.


Tests for q2_1 (points: 1)
```
>>> imported = True
>>> try:
...     from lab03toolkit import multiplicative_inverse
... except ImportError:
...     imported = False
>>> 
>>> imported == True
True
```
failure_message: Your function was not imported correctly
```
>>> multiplicative_inverse(13, 27)
25
```
failure_message: Your function was unable to compute the correct inverse
