```python
# Initialize Otter
import otter
grader = otter.Notebook("Enigma_project.ipynb")
```

# Porta's Disk

The Porta's disk you made enciphers with a version of the Vigenere Cipher. To see how, try enciphering "AAAAAAA" using a key of B2.

One way to program this cipher is to make a string, like this.

```python
string = "BCDEFGHIJKLMNOPQRSTUVWXYZA"
```

Notice that the string now lines up very nicely with our LETTERS. Looking at our wheel and letting the LETTERS represent the plaintext, lowercase letters, of the inner wheel while the string called "string" represents the ciphertext, uppercase letters, of the outer wheel.

```python
string =  "BCDEFGHIJKLMNOPQRSTUVWXYZA" # ciphertext
LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ" # plaintext
```

The trick is that with each letter we encipher or decipher, we need to make the string scoot <i>twice</i> to the left relative to LETTERS and the letter on the left needs to move to the back of the line each time. We can accomplish that with the indexing function that Python does for us.

NOTE: The string scoots twice because the second part of the key is "2"

```python
string = "BCDEFGHIJKLMNOPQRSTUVWXYZA"
key = ("B", 2)

for i in range(key[1]):
    string = string[1:] + string[0]
    print(i + 1, string) # I use "i + 1" here to align more closely with the typical way we think about 
                         # these things- we rotate the wheel once and then twice.
    print(' ', LETTERS.lower())
```

Now, write a function that takes in a message and a key which you can use to decipher this message: "UZKNS JMFSQ UJORA YCOX" You will find text_clean and text_block in the toolkit as usual.

```python
from project_toolkit import *
```

```python
def portasdisk(message, key, encipher = True, LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ"):
    output = ''
    string = LETTERS
    
    ...
    
    if encipher == False:
        return output.lower()
    if encipher == True:
        return text_block(output)

```

```python
grader.check("q1")
```

# Enigma Machine: Part 1

Take a close look at your pringles can enigma machine. In the cell below, I have lined up the letters with the rotors in order. The reason rotor III is in position "D" instead of "C" is because we turn the rotor before beginning to encipher the first letter. The reason is that, in the original machine, the action of pushing down the key provided the force to turn the rotor.

Use these strings to write a program that can correctly encipher/decipher the first message from the activity: "AEFAE  JXXBN  XYJTY"

Notice that the key is hard coded in this program. That's okay for now. Also, note that I have dropped the convention of lower case for plaintext and uppercase for ciphertext. I have also chosen not to block in this part of the assignment. This is to more closely approximate the function of the original enigma machine where such conventions were not possible.

```python
Letters =  "ABCDEFGHIJKLMNOPQRSTUVWXYZ"
RoterIII = "HJLCPRTXVZNYEIWGAKMUSQOBDF" # position D
RoterII =  "OEAJDKSIRUXBLHWTMCQGZNPYFV" # position B 
RoterI =   "JEKMFLGDQVZNTOWYHXUSPAIBRC" # position A
Reflect =  "YRUHQSLDPXNGOKMIEBFZCWVJAT"
```

```python
def enigma(message, Letters =  "ABCDEFGHIJKLMNOPQRSTUVWXYZ"):
    
    RoterIII = "HJLCPRTXVZNYEIWGAKMUSQOBDF" # position D
    RoterII =  "OEAJDKSIRUXBLHWTMCQGZNPYFV" # position B 
    RoterI =   "JEKMFLGDQVZNTOWYHXUSPAIBRC" # position A
    Reflect =  "YRUHQSLDPXNGOKMIEBFZCWVJAT"
    
    output = ''
    for i in message:
        
    ...
  
    return(output)
```

```python
grader.check("q2")
```

# Enigma Machine: Part 2 

One of the things that made the Enigma machine code so difficult to crack was the different rotation frequencies of the different rotors. The middle rotor turns once only every 26 letters. Starting from position RBA on your pringles can, the middle rotor turns after three characters. In order to accomplish this, I have added a counter. The counter starts at position 23 and increments every letter. When it reaches 26, rotor two "turns".

Try writing a function that can decipher this message "MABEK GZXSG".

```python
counter = 23
Letters =  "ABCDEFGHIJKLMNOPQRSTUVWXYZ"
RoterIII = "GAKMUSQOBDFHJLCPRTXVZNYEIW" #position s
RoterII =  "BLHWTMCQGZNPYFVOEAJDKSIRUX" #position b 
RoterI =   "JEKMFLGDQVZNTOWYHXUSPAIBRC" #position A
Reflect =  "YRUHQSLDPXNGOKMIEBFZCWVJAT"
```

```python
def enigma2(message, Letters =  "ABCDEFGHIJKLMNOPQRSTUVWXYZ"):
    output = ""
    
    counter = 23
    
    RoterIII = "GAKMUSQOBDFHJLCPRTXVZNYEIW" #position s
    RoterII =  "BLHWTMCQGZNPYFVOEAJDKSIRUX" #position b 
    RoterI =   "JEKMFLGDQVZNTOWYHXUSPAIBRC" #position A
    Reflect =  "YRUHQSLDPXNGOKMIEBFZCWVJAT"

    ...
    
    return(output)
```

```python
grader.check("q3")
```

# Enigma Machine: Part 3

With rotors I, II, and III in positions A, D, and S (really T) and the following message "RZFOG FYHPL", all three rotors will turn at least once. This time, we need to have two counters. When counter one reaches 26, it will cause rotor 2 to "turn". When counter two reaches 26, it will cause rotor 3 to turn. Watch out! When rotor 3 turns, this will effect the reflector as well- in a way that is not easy to predict! 

<!-- BEGIN QUESTION -->



```python
counter1 = 24
counter2 = 23
letters =  "ABCDEFGHIJKLMNOPQRSTUVWXYZ"
rotorIII = "AKMUSQOBDFHJLCPRTXVZNYEIWGA" # position T
rotorII =  "LHWTMCQGZNPYFVOEAJDKSIRUXB"  # position D 
rotorI =   "RCJEKMFLGDQVZNTOWYHXUSPAIB"  # position A
reflect =  "YRUHQSLDPXNGOKMIEBFZCWVJAT"
```

```python
def enigma3(message):
    output = ""

    ...
        
    return(output)    
    

```

```python
grader.check("q4")
```

<!-- END QUESTION -->

# BONUS

Can you find a way to write function that will take in a message and a key and will automatically set the rotors so that it is consistent with the key? There is no test for this program beyond that which you can perform using your pringles can.

```python
def enigma_epic(message, key):
    output = ''
    
```

# Submitting Your Project

When you finish this project, rename it with a unique part of your name– for example, your first name, if your first name is unique in the class. Otherwise, your first name and first initial of your last name. Then, download this notebook and send it to me attached to an email with the word "project" in the subject line.

In the body of your email, give me at <b>least three</b> messages that I can successfully decipher with your program.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Before exporting, SAVE your notebook, then RESTART AND RUN ALL CELLS. This will run the export cell. Make sure you submit the most recent copy by using the date-time stamp on the file.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 2)
```
>>> portasdisk("It always seems impossible until it's done", ("M", 5)) != None
True
```
failure_message: Check that you are returning values, not printing them
```
>>> portasdisk("UKWMC LONSJ OBMHQ YCLQL JYWRP AUCEU YOEIE", ("M", 5), False)
'italwaysseemsimpossibleuntilitsdone'
```
failure_message: Check that you are returning values, not printing them


Tests for q2 (points: 2)
```
>>> enigma("CUQXRFZCSAX") != None
True
```
failure_message: Check that you are returning values, not printing them
```
>>> enigma("CUQXRFZCSAX")
'ACHIEVEMENT'
```
failure_message: Check that you are returning values, not printing them


Tests for q3 (points: 2)
```
>>> enigma2("PGDOBIIFSBLU") != None
True
```
failure_message: Check that you are returning values, not printing them
```
>>> enigma2("PGDOBIIFSBLU")
'FLYINGCOLORS'
```
failure_message: Check that you are returning values, not printing them


Tests for q4 (points: 2)
```
>>> enigma3("RBFBD") != None
True
```
failure_message: Check that you are returning values, not printing them
```
>>> enigma3("RBFBD")
'THREE'
```
failure_message: Check that you are returning values, not printing them
